# Week 9 Problem Set

## Homeworks

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
```

**HW0.** Copy and paste some of the functions from you cohort sessions and previous exercises that you will need in this homework. See below template and the list here:
- `normalize_z()`
- `get_features_targets()`
- `prepare_feature()`
- `prepare_target()`
- `predict()`
- `predict_norm()`
- `split_data()`
- `r2_score()`
- `mean_squared_error()`

Then do the following:
- Read the CSV file `housing_processed.csv` and extract the following columns:
    - x data: RM, DIS, and INDUS columns 
    - y data: MEDV column
- Normalize the features using z normalization. 
- Plot the data using scatter plot. Use the following columns:


```python
def normalize_z(df):
    ###
    ### YOUR CODE HERE
    ###
    pass

def get_features_targets(df, feature_names, target_names):
    ###
    ### YOUR CODE HERE
    ###
    pass
    return df_feature, df_target

def prepare_feature(df_feature):
    ###
    ### YOUR CODE HERE
    ###
    pass

def prepare_target(df_target):
    ###
    ### YOUR CODE HERE
    ###
    pass

def predict(df_feature, beta):
    ###
    ### YOUR CODE HERE
    ###
    pass

def predict_norm(X, beta):
    ###
    ### YOUR CODE HERE
    ###
    pass

def split_data(df_feature, df_target, random_state=None, test_size=0.5):
    ###
    ### YOUR CODE HERE
    ###
    return df_feature_train, df_feature_test, df_target_train, df_target_test
  
def r2_score(y, ypred):
    ###
    ### YOUR CODE HERE
    ###
    pass

def mean_squared_error(target, pred):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

```python
# Read the CSV file
df = None

# Extract the features and the targets
df_features, df_target = None, None

# Normalize using z normalization
df_features = None

###
### YOUR CODE HERE
###

```

<!-- grade_id: hw01, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw02, points: 1, grade: true, solution: false -->
```python
display(df_features.describe())
display(df_target.describe())
assert np.isclose(df_features['RM'].mean(), 0)
assert np.isclose(df_features['DIS'].mean(), 0)
assert np.isclose(df_features['INDUS'].mean(), 0)

assert np.isclose(df_features['RM'].std(), 1)
assert np.isclose(df_features['DIS'].std(), 1)
assert np.isclose(df_features['INDUS'].std(), 1)

assert np.isclose(df_target['MEDV'].mean(), 22.532806)
assert np.isclose(df_target['MEDV'].std(), 9.197104)

assert np.isclose(df_features['RM'].median(), -0.1083583)
assert np.isclose(df_features['DIS'].median(), -0.2790473)
assert np.isclose(df_features['INDUS'].median(), -0.2108898)
```

<!-- grade_id: hw03, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
sns.set()
plt.scatter(df_features["RM"], df_target)
```

```python
plt.scatter(df_features["DIS"], df_target)
```

```python
plt.scatter(df_features["INDUS"], df_target)
```

**HW1.** *Multiple variables cost function:* Write a function `compute_cost_multi()` to compute the cost function of a linear regression model. The function should take in two 2-D numpy arrays. The first one is the matrix of the linear equation and the second one is the actual target value.

Recall that:

$$J(\hat{\beta}_0, \hat{\beta}_1) = \frac{1}{2m}\Sigma^m_{i=1}\left(\hat{y}(x^i)-y^i\right)^2$$

where

$$\hat{y}(x) =  \hat{\beta}_0 + \hat{\beta}_1 x_1 + \hat{\beta}_2 x_2 + \ldots + \hat{\beta}_n x_n$$

The function should receive three Numpy arrays:
- `X`: is the feature 2D Numpy array
- `y`: is the target 2D Numpy array
- `beta`: is the parameter 2D Numpy array

The function should return the cost which is a float.

You can use the following function in your code:
- `np.matmul(array1, array2)`

Note that if you wrote your Cohort session's `compute_cost()` using proper Matrix operations to do the square and the summation, the code will be exactly the same here.

```python
def compute_cost(X, y, beta):
    J = 0
    ###
    ### YOUR CODE HERE
    ###
    return J

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
X = prepare_feature(df_features)
target = prepare_target(df_target)

assert isinstance(X, np.ndarray)
assert isinstance(target, np.ndarray)
assert X.shape == (506, 4)
assert target.shape == (506, 1)
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
beta = np.zeros((4,1))
J = compute_cost(X, target, beta)
print(J)
assert np.isclose(J, 296.0734)

beta = np.ones((4,1))
J = compute_cost(X, target, beta)
print(J)
assert np.isclose(J, 270.4083)

beta = np.array([-1, 2, 1, 2]).reshape((4,1))
J = compute_cost(X, target, beta)
print(J)
assert np.isclose(J, 314.8510)

```

<!-- grade_id: hw14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Gradient Descent:* Write a function called `gradient_descent_multi()` that takes in four parameters:
- `X`: is a 2-D numpy array for the features
- `y`: is a vector array for the target
- `alpha`: is the learning rate
- `num_iters`: is the number of iteration to perform

The function should return two arrays:
- `beta`: is coefficient at the end of the iteration
- `J_storage`: is the array that stores the cost value at each iteration

You can use some of the following functions:
- `np.matmul(array1, array2)`: which is to do matrix multiplication on two Numpy arrays.
- `compute_cost()`: which the function you created in the previous problem set to compute the cost.

Note that if you use proper matrix operations in your cohort sessions for the gradient descent function, the code will be the same here.

```python
def gradient_descent(X, y, beta, alpha, num_iters):
    ###
    ### YOUR CODE HERE
    ###
    return beta, J_storage

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
iterations = 1500
alpha = 0.01
beta = np.zeros((4,1))

beta, J_storage = gradient_descent(X, target, beta, alpha, iterations)
print(beta)
assert np.isclose(beta[0], 22.5328)
assert np.isclose(beta[1], 5.4239)
assert np.isclose(beta[2], -0.90367)
assert np.isclose(beta[3], -2.95818)
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
plt.plot(J_storage)
```

**HW3.** Do the following tasks:
- Get the features and the targets.
    - features: RM, DIS, INDUS
    - target: MEDV
- Use the previous functions called `predict()` to calculated the predicted values given the features and the model. 
- Create a target numpy array from the data frame.

```python
# This is from the previous result
beta = np.array([[22.53279993],
       [ 5.42386374],
       [-0.90367449],
       [-2.95818095]])

# Extract the feature and the target
df_features, df_target = None, None

# Call predict()
pred = None

# Change target to numpy array
target = None

###
### YOUR CODE HERE
###


```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert isinstance(pred, np.ndarray)
assert np.isclose(pred.mean(), 22.5328)
assert np.isclose(pred.std(), 6.7577)
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
plt.scatter(df_features["RM"],target)
plt.scatter(df_features["RM"],pred)
### END HIDDEN TESTS
```

```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *Splitting data:*  Do the following tasks:
- Extract the following:
    - features: RM, DIS, and INDUS
    - target: MEDV 
- Use Week 9's function `split_data()` to split the data into train and test using `random_state=100` and `test_size=0.3`.
- Normalize and prepare the features and the target.
- Use the training data set and call `gradient_descent()` to obtain the `theta`.
- Use the test data set to get the predicted values.

You need to replace the `None` in the code below with other a function call or any other Python expressions. 

```python
# Extract the features and the target
df_features, df_target = None, None

# Split the data set into training and test
df_features_train, df_features_test, df_target_train, df_target_test = None, None, None, None

# Normalize the features using z normalization
df_features_train_z = None

# Change the features and the target to numpy array using the prepare functions
X = None
target = None

iterations = 1500
alpha = 0.01
beta = np.zeros((4,1))

# Call the gradient_descent function
beta, J_storage = None, None

# call the predict() method
pred = None

###
### YOUR CODE HERE
###

```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
assert isinstance(pred, np.ndarray)
assert pred.shape == (151, 1)
assert np.isclose(pred.mean(), 22.66816)
assert np.isclose(pred.std(), 6.67324)
```

<!-- grade_id: hw43, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
plt.scatter(df_features_test["RM"], df_target_test)
plt.scatter(df_features_test["RM"], pred)
```

```python
plt.scatter(df_features_test["DIS"], df_target_test)
plt.scatter(df_features_test["DIS"], pred)
```

```python
plt.scatter(df_features_test["INDUS"], df_target_test)
plt.scatter(df_features_test["INDUS"], pred)
```

**HW5.** Calculate the coefficient of determination, $r^2$.

```python
# change target test set to a numpy array
target = None

# Calculate r2 score by calling a function
r2 = None

###
### YOUR CODE HERE
###

print(r2)

```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
assert np.isclose(r2, 0.47713)
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW6.** Calculate the mean squared error.

```python
# Calculate the mse
mse = None

###
### YOUR CODE HERE
###
print(mse)

```

<!-- grade_id: hw61, points: 1, grade: true, solution: false -->
```python
assert np.isclose(mse, 51.363)
```

<!-- grade_id: hw62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW7.** *Polynomial Transformation:* Redo the steps for breast cancer data but this time we will use quadratic model. Use the following columns:
- x data: radius_mean
- y data: area_mean

We will create a quadratic hypothesis for this x and y data. To do that write a function `transform_features(df, colname, colname_transformed)` that takes in a dataframe for the features, the original column name, and the transformed column name. The function should add another column with the name `colname_transformed` with the value of column in `colname` transformed to its quadratic value. 

```python
# Read from breast_cancer_data.csv file
df = None

# Extract feature and target
df_feature, df_target = None, None

###
### YOUR CODE HERE
###

plt.scatter(df_feature, df_target)
```

<!-- grade_id: hw71, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# write your function to create a quadratic feature of x

def transform_features(df_feature, colname, colname_transformed):
    ###
    ### YOUR CODE HERE
    ###
    return df_feature

```

<!-- grade_id: hw72, points: 1, grade: true, solution: false -->
```python
df_features = transform_features(df_feature, "radius_mean", "radius_mean^2")

assert (df_features.loc[:,"radius_mean^2"] == df_features.loc[:,"radius_mean"] ** 2).all()
```

<!-- grade_id: hw73, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# split data using random_state = 100 and 30% test size
df_features_train, df_features_test, df_target_train, df_target_test = None, None, None, None

# normalize features
df_features_train_z = None

# change to numpy array and append column for feature
X = None
target = None

iterations = 1500
alpha = 0.01
beta = np.zeros((3,1))

# call gradient_descent() function
beta, J_storage = None, None

###
### YOUR CODE HERE
###

```

<!-- grade_id: hw74, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw75, points: 1, grade: true, solution: false -->
```python

assert np.isclose(beta[0], 646.0787)
assert np.isclose(beta[1], 146.4801)
assert np.isclose(beta[2], 201.9803)
```

<!-- grade_id: hw76, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
plt.plot(J_storage)
```

```python
# change target to numpy array
beta = np.array([[646.0787641 ], [146.4800792 ], [201.98031254]])

target = None
pred = None

# get predicted values
###
### YOUR CODE HERE
###

```

<!-- grade_id: hw77, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
plt.scatter(df_features_test["radius_mean"], target)
plt.scatter(df_features_test["radius_mean"], pred)
```

<!-- grade_id: hw78, points: 1, grade: true, solution: false -->
```python
r2 = r2_score(df_target_test, pred)
print(r2)
assert np.isclose(r2, 0.985095)
```

<!-- grade_id: hw79, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw710, points: 1, grade: true, solution: false -->
```python
mse = mean_squared_error(df_target_test, pred)
print(mse)
assert np.isclose(mse, 1919.164)
```

<!-- grade_id: hw711, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW8.** *Optional:* Redo the above tasks using Sci-kit learn libraries. You will need to use the following:
- [LinearRegression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html)
- [train_test_split](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html)
- [r2_score](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.r2_score.html)
- [mean_squared_error](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.mean_squared_error.html)
- [PolynomialFeatures](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.PolynomialFeatures.html)

```python
from sklearn.linear_model import LinearRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import r2_score, mean_squared_error
from sklearn.preprocessing import PolynomialFeatures
```

**Redo HW 4 using Scikit Learn**

```python
# Read the housing_processed.csv file
df = None

# extract the features from ["RM", "DIS", "INDUS"] and target from []"MEDV"]
df_features, df_target = None, None
# normalize
df_features = None

###
### YOUR CODE HERE
###

```

```python
# Split the data into training and test data set using scikit-learn function
df_features_train, df_features_test, df_target_train, df_target_test = None, None, None, None

# Instantiate LinearRegression() object
model = None

# Call the fit() method
pass

###
### YOUR CODE HERE
###
print(model.coef_, model.intercept_)
assert np.isclose(model.coef_, [ 5.01417104, -1.00878266, -3.27301726]).all()
assert np.isclose(model.intercept_, 22.45962454)
```

```python
# Call the predict() method
pred = None

###
### YOUR CODE HERE
###

```

```python
plt.scatter(df_features_test["RM"], df_target_test)
plt.scatter(df_features_test["RM"], pred)
```

```python
plt.scatter(df_features_test["DIS"], df_target_test)
plt.scatter(df_features_test["DIS"], pred)
```

```python
plt.scatter(df_features_test["INDUS"], df_target_test)
plt.scatter(df_features_test["INDUS"], pred)
```

```python
r2 = r2_score(df_target_test, pred)
print(r2)
assert np.isclose(r2, 0.48250)
```

```python
mse = mean_squared_error(df_target_test, pred)
print(mse)
assert np.isclose(mse, 52.41451)
```

**Redo HW7 Using Scikit Learn**

```python
# Read the file breast_cancer_data.csv
df = None
# extract feature from "radius_mean" and target from "area_mean"
df_feature, df_target = None, None

###
### YOUR CODE HERE
###

```

```python
# instantiate a PolynomialFeatures object with degree = 2
poly = None

# call its fit_transform() method
df_features = None

# call train_test_split() to split the data
df_features_train, df_features_test, df_target_train, df_target_test = None, None, None, None

# instantiate LinearRegression() object
model = None

# call its fit() method
pass


### BEGIN SOLUTON
poly = PolynomialFeatures(2)
df_features = poly.fit_transform(df_feature)
df_features_train, df_features_test, df_target_train, df_target_test = train_test_split(df_features, df_target, random_state=100, test_size=0.3)
model = LinearRegression()
model.fit(df_features_train, df_target_train)
### END SOLUTION
print(model.coef_, model.intercept_)
assert np.isclose(model.coef_, [0., 3.69735512, 2.9925278 ]).all()
assert np.isclose(model.intercept_, -32.3684598)
```

```python
# Call the predict() method
pred = None

###
### YOUR CODE HERE
###

print(type(pred), pred.mean(), pred.std())
assert isinstance(pred, np.ndarray)
assert np.isclose(pred.mean(), 672.508465)
assert np.isclose(pred.std(), 351.50271)
```

```python
plt.scatter(df_features_test[:,1], df_target_test)
plt.scatter(df_features_test[:,1], pred)
```

```python
r2 = r2_score(df_target_test, pred)
print(r2)
assert np.isclose(r2, 0.99729)
```

```python
mse = mean_squared_error(df_target_test, pred)
print(mse)
assert np.isclose(mse, 346.79479)
```