# Week 6 Problem Set

## Homeworks

**HW1.** Extend the class `Fraction` to implement the other operators: `- * < <= > >=`.

```python
def gcd(a, b):
    if b == 0:
        return a
    else:
        return gcd(b, a % b)


class Fraction:
    # copy the rest of the methods here
    
    def __init__(self, num, den):
        self.num = num
        self.den = den
    
    @property
    def num(self):
        return self.numerator
    
    @num.setter
    def num(self, val):
        self.numerator = int(val)
    
    @property
    def den(self):
        return self.denominator
    
    @den.setter
    def den(self, val):
        if val != 0:
            self.denominator = int(val)
        else:
            self.denominator = 1
    
    def __str__(self):
        return f"{self.numerator}/{self.denominator}"

    def simplify(self):
        common_den = gcd(self.numerator, self.denominator)
        num = self.numerator // common_den
        den = self.denominator // common_den
        return Fraction(num, den)
        
    
    def __add__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator 
        
        numerator = a_num * b_den + b_num * a_den
        denominator = a_den * b_den
        result = Fraction(numerator, denominator)
        return result.simplify()
        
        
    def __sub__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        numerator = a_num * b_den - b_num * a_den
        denominator = a_den * b_den
        result = Fraction(numerator, denominator)
        return result.simplify()
        
    
    def __mul__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        numerator = a_num * b_num
        denominator = a_den * b_den
        result = Fraction(numerator, denominator)
        return result.simplify()
    
    def __eq__(self, other):
        left = self.simplify()
        right = other.simplify()
        return (left.numerator == right.numerator) and (left.denominator == right.denominator)
    
    
    def __lt__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        return (a_num / a_den) < (b_num / b_den)
    
    def __le__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        return (a_num / a_den) <= (b_num / b_den)
    
    def __gt__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        return (a_num / a_den) > (b_num / b_den)
    
    def __ge__(self, other):
        a_num = self.numerator
        a_den = self.denominator
        b_num = other.numerator
        b_den = other.denominator
        
        return (a_num / a_den) >= (b_num / b_den)
    

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
f1 = Fraction(3, 4)
f2 = Fraction(1, 2)
f3 = f1 - f2
assert f3 == Fraction(1, 4)
f4 = f1 * f2
assert f4 == Fraction(3, 8)
assert f2 < f1
assert f2 <= f2
assert f1 > f3
assert f3 >= f3
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# Copy your solution from the Cohort problem
class MixedFraction(Fraction):
    
    def __init__(self, top, bot, whole=0):
        num = whole * bot + top
        super().__init__(num, bot) #calling parent class

    def get_three_numbers(self):
        whole = self.num // self.den
        top = self.num % self.den
        bot = self.den
        return (top, bot, whole)

    def __str__(self):
        (top, bot, whole) = self.get_three_numbers()
        if whole > 0:
            return f"{whole} {top}/{bot}"
        else:
            return f"{top}/{bot}"

```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
assert result.num == 10 and result.den == 3

result = mf1 * mf2
assert result.num == 25 and result.den == 9

mf3 = MixedFraction(1, 2, 1)
result = mf1 - mf3
assert result.num == 1 and result.den == 6

assert str(mf1) == "1 2/3"
```

<!-- grade_id: hw14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** Write a class called `EvaluateFraction` that evaluates postfix notation implemented using Dequeue data structures only. Postfix notation is a way of writing expressions without using parenthesis. For example, the expression `(1+2)*3` would be written as `1 2 + 3 *`. The class `EvaluateFraction` has the following method:
- `input(inp)`: which pushes the input input one at a time. For example, to create a postfix notation `1 2 + 3 *`, we can call this method repetitively, e.g. `e.input('1'); e.input('2'); e.input('+'); e.input('3'); e.input('*')`. Notice that the input is of String data type. 
- `evaluate()`: which returns the output of the expression.
- `get_fraction(inp)`: which takes in an input string and returns a `Fraction` object. 

Postfix notation is evaluated using a Stack. Since `Dequeue` can be used for both Stack and Queue, we will implement using `Dequeue`. The input streams from `input()` are stored in a Queue, which we will again implement using Dequeue. If the output of the Queue is a number, the item is pushed into the stack. If it is an operator, we will apply the operator to the two top most item n the stacks and push the result back into the stack. 

```python
class Stack:
    def __init__(self):
        self.__items = []
        
    def push(self, item):
        self.__items.append(item)

    def pop(self):
        if not self.is_empty:
            return self.__items.pop()

    def peek(self):
        if not self.is_empty:
            return self.__items[-1]

    @property
    def is_empty(self):
        return self.__items == []

    @property
    def size(self):
        return len(self.__items)

```

```python
class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()

    @property
    def is_empty(self):
        return self.left_stack.is_empty and self.right_stack.is_empty

    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size

    def enqueue(self, item):
        self.right_stack.push(item)
    
    def right_to_left (self):
        while not self.right_stack.is_empty:
            item = self.right_stack.pop()
            self.left_stack.push(item)
    
    def dequeue(self):
        if self.left_stack.is_empty:
            self.right_to_left()
        return self.left_stack.pop()
    
    def peek(self):
        if self.left_stack.is_empty:
            self.right_to_left()
        return self.left_stack.peek()
```

```python
class Dequeue(Queue):
  
    def add_front(self, item):
        self.left_stack.push(item)
      
    def remove_front(self):
        return self.dequeue()
    
    def add_rear(self, item):
        self.enqueue(item)
    
    def remove_rear(self):
        if self.right_stack.is_empty:
            self.left_to_right()
        return self.right_stack.pop()
    
    def left_to_right (self):
        while not self.left_stack.is_empty:
            item = self.left_stack.pop()
            self.right_stack.push(item)
    
    def peek_front(self):
        return self.peek()
    
    def peek_rear(self):
        if self.right_stack.is_empty:
            self.left_to_right()
        return self.right_stack.peek()
```

```python
class EvaluateFraction:

    operands = "0123456789"
    operators = "+-*/"
    
    def __init__(self):
        self.expression = Dequeue()
        self.stack = Dequeue()
    
    def input(self, item):
        if item in self.operators:
            self.stack.add_front(item)
        else:
            self.expression.add_front(item)
    
    def evaluate(self):
        while not self.stack.is_empty:
            item = self.stack.remove_rear()
            if item in self.operators and len(item) == 1:
                left = self.get_fraction(self.expression.remove_rear())
                right = self.get_fraction(self.expression.remove_rear())
                result = self.process_operator(left, right, item)
                self.expression.add_rear(result)
            else:
                self.expression.add_rear(int(item))
        return self.expression.remove_rear()
    
    def get_fraction(self, inp):
        _input = str(inp)
        fract = _input.split('/')
        num = fract[0]
        den = fract[1]
        return Fraction(num, den)
    
    def process_operator(self, op1, op2, op):
        if op == "+":
            return op1 + op2
        elif op == "-":
            return op1 - op2
        elif op == "*":
            return op1 * op2
        elif op == "/":
            return op1 / op2

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
pe = EvaluateFraction()
pe.input("1/2")
pe.input("2/3")
pe.input("+")
assert pe.evaluate()==Fraction(7, 6)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
assert pe.evaluate()==Fraction(1, 1)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
pe.input("3/4")
pe.input("*")
assert pe.evaluate()==Fraction(3, 4)
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Modify HW2 so that it can work with MixedFraction. Write a class called `EvaluateMixedFraction` as a subclass of `EvaluateFraction`. You need to override the following methods:
- `get_fraction(inp)`: This function should be able to handle string input for MixedFraction such as `1 1/2` or `3/2`. It should return a `MixedFraction` object.
- `evaluate()`: This function should return `MixedFraction` object rather than `Fraction` object. 

```python
class EvaluateMixedFraction(EvaluateFraction):
    def get_fraction(self, inp):
        whole = 0
        _input = str(inp)
        mixed = _input.split(' ')
        if len(mixed) > 1:
            whole = mixed[0]
            fract = mixed[1]
            fract = fract.split('/')
            num = fract[0]
            den = fract[1]
        else:
            fract = _input.split('/')
            num = fract[0]
            den = fract[1]
        return MixedFraction(int(num), int(den), int(whole))
    
    def evaluate(self):
        answer = None
        while not self.stack.is_empty:
            item = self.stack.remove_rear()
            if item in self.operators and len(item) == 1:
                left = self.get_fraction(self.expression.remove_rear())
                right = self.get_fraction(self.expression.remove_rear())
                result = self.process_operator(left, right, item)
                self.expression.add_rear(result)
            else:
                self.expression.add_rear(int(item))
        return self.expression.remove_rear()
    
        return MixedFraction(answer.num, answer.den)
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
pe = EvaluateMixedFraction()
pe.input("3/2")
pe.input("1 2/3")
pe.input("+")
assert pe.evaluate() == MixedFraction(1, 6, 3)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1 1/8")
pe.input("-")
assert pe.evaluate() == MixedFraction(1, 24)

pe.input("1 1/2")
pe.input("2 2/3")
pe.input("+")
pe.input("1 1/6")
pe.input("-")
pe.input("5/4")
pe.input("*")
assert pe.evaluate() == MixedFraction( 3, 4, 3)
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### YOUR CODE HERE
###

```

**HW4.** *Linked List:* We are going to implement Linked List Abstract Data Type. To do so, we will implement two classes: `Node` and `MyLinkedList`. In this part, we will implement the class Node.

The class `Node` has the following attribute and computed property:
- `element`: which stores the value of the item in that node.
- `next`: which stores the reference to the next `Node` in the list. The setter method should check if the value assigned is of type `Node`.





```python
class Node:
    def __init__(self, e):
        self.element = e
        self.__next = None
               
    @property
    def next(self):
        return self.__next
    
    @next.setter
    def next(self, value):
        # check if value is an instance of Node object
        # you can use isinstance() function
        if isinstance(value, Node):
            self.__next = value

```

**HW5.** This is a continuation to implement a Linked List. The class `MyLinkedList` has two different properties:
- `head`: which points to the `Node` of the first element.
- `tail`: which points to the `Node` of the last element.

It should also have the following methods:
- `__init__(items)`: which create the link list object based using the arguments.
- `get(index)`: which returns the item at the given `index`.
- `add_first(item)`: which adds the `item` as the first element.
- `add_last(item)`: which adds the `item` as the last element.
- `add_at(index, item)`: which adds the `item` at the position `index`.
- `remove_first(item)`: which removes the `item` as the first element.
- `remove_last(item)`: which removes the `item` as the last element.
- `remove_at(index, item)`: which removes the `item` at the position `index`.


```python
import collections.abc as c

class MyAbstractList(c.Iterator):
    
    
    def __init__(self, list_items):
        # iterate over every element and call self.add(item)
        self.size = 0
        self._idx = 0
        
        for items in list_items:
            self.append(items)
    
    @property
    def is_empty(self):
        return self.size == 0
    
    def append(self, item):
        # call add_at() method here
        self.add_at(self.size, item)
        
    def remove(self, item):
        # you should use remove_at() method here
        index = self.index_of(item)
        if index >= 0:
            self.remove_at(index)
        
    def __getitem__(self, index):
        return self.get(index)
    
    def __setitem__(self, index, value):
        # call set_at(index, value) method here
        self.set_at(index, value)
        
    def __delitem__(self, index):
        # call remove_at() method here
        self.remove_at(index)
    
    def __len__(self):
        return self.size
        
    def __iter__(self):
        self._idx = 0
        return self
        
    def __next__(self):
        if self._idx < self.size:
            n_item = self.get(self._idx)
            self._idx += 1
            return n_item
        else:
            raise StopIteration
```

```python
class MyLinkedList(MyAbstractList):
    def __init__(self, items):
        self.head = Node(items[0])
        self.tail = Node(items[-1])
        super().__init__(items)
        
    def get(self, index):
        # do the following:
        # 1. traverse to the node at index
        # 2. return the element of that node
        current = self.head
        iter_idx = 0
        while current != None:
            if iter_idx == index:
                return current.element
            
            iter_idx += 1
            current = current.next
            
    
    def add_first(self, element):
        # do the following:
        # 1. create a new Node object using element
        # 2. set the current head reference as the next reference of the new node
        # 3. increase size by 1
        # 4. if this is the last element (no tail) -> set the current node as the tail
        new_node = Node(element)
        
        new_node.next = self.head
        self.head = new_node
        self.size += 1
        if self.size == 1:
            self.tail = self.head
        
            
    def add_last(self, element):
        # do the following:
        # 1. create a new Node object using element
        # 2. if there is no element as tail -> set the new node as both
        #    the tail and the head
        # 3. otherwise, -> 
        #    - set the new node as the next reference of the tail
        #    - set the next reference of the current node as the tail's next reference
        # 4. increase size by 1
        
        new_node = Node(element)
        
        if self.head is None:
            self.head = new_node
            self.tail = new_node
        else:
            current_node = self.tail
            self.tail = new_node
            current_node.next = self.tail
            
        self.size += 1
            
        
    def add_at(self, index, element):
        if index == 0:
            # if insert at first position, call add_first() method
            self.add_first(element)
            
        elif index >= self.size:
            # if insert at last position, call add_last() method
            self.add_last(element)
        else:
            # if insert in between, do the following:
            # 1. start from the head, traverse the linked list to get
            #    the reference at position index using its next reference
            # 2. create a new Node
            # 3. set the next of the current node as the next of the new Node
            # 4. set the new node as the next of the current node
            # 5. increase the size by 1
            current = self.head
            iter_idx = 1
            while current.next != None:
                if iter_idx == index:
                    new_node = Node(element)
                    new_node.next = current.next
                    current.next = new_node
                iter_idx += 1
                current = current.next
            
            self.size += 1
            
            
            
    def set_at(self, index, element):
        if 0 <= index < self.size:
            current = self.head
            for idx in range(0, index):
                current = current.next
            current.element = element
            
    def remove_first(self):
        if self.size == 0:
            # if list is empty, return None
            return None
        else:
            # otherwise, do the following:
            # 1. store the head at a temporary variable
            # 2. set the next reference of the current head to be the head
            # 3. reduce size by 1
            # 4. if the new head is now None, it means empty list
            #    -> set the tail to be None also
            # 5. return element of the removed node
            temp = self.head
            self.head = self.head.next
            self.size -= 1
            
            if self.head == None:
                self.tail = None
                
            return temp
        
    def remove_last(self):
        if self.size == 0:
            # if the list is empty, return None
            return None
        elif self.size == 1:
            # if there is only one element, just remove that one node 
            # using some other method
            self.remove_first()
        else:
            # otherwise, do the following:
            # 1. traverse to the second last node
            # 2. store the tail of the list to a temporary variable
            # 3. set the current node as the tail
            # 4. set the next ref of the tail to be None
            # 5. reduce the size by 1
            # 6. return the element of the removed node in the temp var
            current = self.head
            
            while current.next != self.tail:
                current = current.next
                
            temp = self.tail
            self.tail = current
            current.next = None
            self.size -= 1
            return temp
        
    def remove_at(self, index):
        if index < 0 or index >= self.size:
            return None
        elif index == 0:
            return self.remove_first()
        elif index == self.size - 1:
            return self.remove_last()
        else:
            # do the following:
            # 1. traverse to the node at index - 1
            # 2. get the node at index using next reference
            # 3. set the next node of the node at index - 1
            # 4. decrease the size by 1
            # 5. return the element that is removed
            current = self.head
            iter_idx = 0
            while iter_idx < (index - 1):
                current = current.next
                iter_idx += 1
                
            node_at_index = current.next
            temp = node_at_index
            current.next = node_at_index.next
            
            self.size -= 1
            return temp
        

```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
asean = MyLinkedList(['Singapore', 'Malaysia'])
assert asean.head.element == 'Singapore'
assert asean.tail.element == 'Malaysia'

asean.append('Indonesia')
assert asean.tail.element == 'Indonesia'
asean.add_at(0, 'Brunei')
assert asean.head.element == 'Brunei'
assert asean.size == 4
assert len(asean) == 4

asean[0] = 'Cambodia'
assert asean[0] == 'Cambodia' and asean[1] == 'Singapore'
asean[2] = 'Myanmar'
assert(len(asean)) == 4 
assert [x for x in asean] == ['Cambodia', 'Singapore', 'Myanmar', 'Indonesia']

del asean[0]
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Indonesia']

asean.add_at(2, 'Brunei')
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Brunei', 'Indonesia']
del asean[3]
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Brunei']
del asean[1]
assert [x for x in asean] == ['Singapore', 'Brunei']
del asean[1]
assert [x for x in asean] == ['Singapore']
del asean[0]
assert [x for x in asean] == []
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```

```python

```