# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

User recursion for your implementation.

```python
def fibonacci(index):
    if index == 0:
        return 0
    elif index == 1:
        return 1
    else:
        return fibonacci(index - 1) + fibonacci(index - 2)

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index):
    return (index * 2) + 1

def right_of(index):
    return (index + 1) * 2

def max_heapify(array, index, heap_size):
    left = left_of(index)
    right = right_of(index)
    if left > heap_size:
        return array
    else:
        if right >= heap_size:
            max_child = left
        elif array[right] > array[left]:
            max_child = right
        else:
            max_child = left
            
        if array[max_child] > array[index]:
            array[max_child], array[index] = array[index], array[max_child]
            max_heapify(array, max_child, heap_size)
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
# def permutate(s):
#     return permutate_recursive("", s)
    
# def permutate_recursive(str1, str2):
    
#     #base case
#     output = []
#     n = len(str2)
#     if n == 0:
#         output.append(str1)
    
#     #recursive case
#     for i in range(len(str2)):
#         str1 = str1 + str2[i]
#         str2 = str2[0:i] + str2[i+1:]
#         permutate_recursive(str1, str2)
#     return output
        
def permutate(s):
    output = []
    result = permutate_recursive(output,"", s)
    return output

def permutate_recursive(output, str1, str2):
    n = len(str2)
    if n == 0:
        output.append(str1)
        
    for i in range(len(str2)):
        tempStr1 = str1 + str2[i]
        tempStr2 = str2[0:i] + str2[i+1:]
        permutate_recursive(output, tempStr1, tempStr2)
    return output    

print(permutate("abc"))
print(permutate("abcd"))

# def permutate(s): #calls the second function at the beginning
#     newS = ""
#     permed = permutater(newS,s)
#     string = ""
#     outlist = []
#     for index in range(0, len(permed), len(s)):
#         #outlist[index] = permed[index:index + len(s)]
#         string = string + '"' + permed[index:index + len(s)] + '"' + ", "
#     string = "[" + string[0:-2] + "]"
#     print(string)
#     return string
#     #return outlist
#     pass


# def permutater(newS, s):
#     TempnewS = ""
#     if len(s) == 1:
#         return newS + s
#     else:
#         for item in s:
#             TempnewS = TempnewS + permutater(newS + item, s.replace(item,''))
#         return TempnewS
```

Output:
```
File "<ipython-input-29-b26c6adda1dc>", line 20
    result = permutate_recursive(output = [],"", s)
                                             ^
SyntaxError: positional argument follows keyword argument
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    if a == b or b == 0:
        return a
    elif a == 0:
        return b
    else:
        if a > b:
            a = a - b
            return gcd(a, b)
        else:
            b = b - a
            return gcd(a, b)
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```