# Week 5 Problem Set

## Cohort Sessions

**CS1.** *Dictionary:* Implement a Graph using a *Dictionary* where the keys are the Vertices in the Graph and the values (in the the key-value pair) correspond to an Array containing the neighbouring Vertices. For example, let's represent the following graph:
```
    A -> B
    A -> C
    B -> C
    B -> D
    C -> D
    D -> C
    E -> F
    F -> C
```
Create a Dictionary to represent the graph above.

```python
# replace the None with a dictionary representing the graph
graph = {'A': ['B', 'C'], 
         'B': ['C', 'D'], 
         'C': ['D'], 
         'D': ['C'], 
         'E': ['F'], 
         'F': ['C']}

```

```python
print(graph)
```

Output:
```
{'A': ['B', 'C'], 'B': ['C', 'D'], 'C': ['D'], 'D': ['C'], 'E': ['F'], 'F': ['C']}
```

Write a function `get_neighbours(graph, vert)` which returns a list of all neighbours of the requested Vertex `vert` in the `graph`. Return `None` if the Vertex is not in the graph.

```python
def get_neighbours(graph, vert):
    return graph.get(vert)
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert get_neighbours(graph, "B") == ["C", "D"]
assert get_neighbours(graph, "A") == ["B", "C"]
assert get_neighbours(graph, "F") == ["C"]
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Write a function `get_source(graph, vert)` which returns a list of all source Vertices pointing to `vert` in the `graph`. For example, Vertex "C" has the following source Vertices: `["A", "B", "D", "F"]. Return an empty list if there are none.

```python
def get_source(graph, vert):
    result = []
    for k in graph:
        for i in graph[k]:
            if i == vert:
                result.append(k)
    return result
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert sorted(get_source(graph, "C")) == ["A", "B", "D", "F"]
assert sorted(get_source(graph, "D")) == ["B", "C"]
assert sorted(get_source(graph, "F")) == ["E"]
```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Create a class `Vertex` to represent a vertex in a graph. The class `Vertex` has the following attributes:
- `id`: to identify each vertex. This is of String data type.
- `neighbours`: which is a Dictionary where the keys are the neighbouring `Vertex` object instances that are connected to the current Vertex and the values are the weights of the edge between the current Vertex and the neighbouring vertices. 

The class should also have the following methods:

- `__init__(self, id)`: which is used to initialized the attribute `id`. By default, `id` is set to an empty String . The attribute `neighbours` is always set to an empty dictionary.
- `add_neighbour(self, nbr_vertex, weight)`: which adds a neighbouring Vertex to the current Vertex. The second argument provides the weight of the edge between the current Vertex and the newly added neighbouring Vertex. By default, `weight` is `0`.
- `get_neigbours(self)`: which returns all the Vertices connected to the current Vertex as a list. The elements of the output list are of `Vertex` object instances.
- `get_weight(self, neighbour)`: which returns the weight of the requested neighbour. It should return `None` if the requested neighbour is not found.
- `__eq__(self, other)`: which returns true if the id of the current vertex object is the same as the `other` vertex's id. 
- `__lt__(self, other)`: which returns true if the id of the current vertex object is less than the `other` vertex's id.
- `__hash__(self)`: which calls the `hash()` function on `id` and returns it. This allows the object to be a dictionary key.
- `__str__(self)`: This method should return the id of the current vertex and a list of `id`s of the neighbouring vertices, like `Vertex 2 is connected to: 3, 4, 5.`

```python
class Vertex:
    def __init__(self, id=""):
        self.id = id
        self.neighbours = {}
    
    def add_neighbour(self, nbr_vertex, weight=0):
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self):
        return list(self.neighbours.keys())
    
    def get_weight(self, neighbour):
        return self.neighbours.get(neighbour)
    
    def __eq__(self, other):
        return self.id == other.id
    
    def __lt__(self, other):
        return self.id < other.id
    
    def __hash__(self):
        return hash(self.id)
    
    def __str__(self):
        return f"Vertex {self.id} is connected to: " + ', '.join([x.id for x in self.get_neighbours()])
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
v1 = Vertex("1")
assert v1.id == "1" and len(v1.neighbours) == 0
v2 = Vertex("2")
v1.add_neighbour(v2)
assert v1.get_neighbours()[0].id == "2" and v1.neighbours[v1.get_neighbours()[0]] == 0
v3 = Vertex("3")
v1.add_neighbour(v3, 3)
assert v1.get_weight(v3) == 3
v4 = Vertex("4")
assert v1.get_weight(v4) == None
assert v1 < v2
assert v1 != v2
assert str(v1) == "Vertex 1 is connected to: 2, 3"
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** Create a class `Graph` to represent a Graph. The class has the following attributes:
- `vertices`: which is a *dictionary* of Vertices. The keys are the `id`s of the Vertices and the values are `Vertex` object instances.

The class has the follo
- `num_vertices`: which is a *computed* property that returns the number of vertices in the graph.

The class also has the following methods:
- `__init__(self)`: which initializes the graph with an empty dictionary.
- `_create_vertex(self, id)`: which creates a new `Vertex` object with a given `id`. This method is never called directly and is only used by `add_vertex(id)`.
- `add_vertex(self, id)`: which creates a new `Vertex` object, adding it into the dictionary `vertices`. The argument `id` is a String. This method should call `_create_vertex(id)`.
- `get_vertex(self, id)`: which returns the `Vertex` object instance of the requested `id`. The method should return `None` if the requested `id` cannot be found. The argument `id` is a String.
- `add_edge(start_v, end_v)`: which creates an edge from one Vertex to another Vertex. The arguments are the `id`s of the two vertices and are both Strings.
- `get_neighbours(self, id)`: which returns a list of `id`s all the neighbouring vertices (of the specified Vertex `id`). It should return `None` if `id` cannot be found. The argument `id` is a String and the elements of the output list are of `str` data type. 
- `__contains__(self, id)`: which returns either `True` or `False` depending on whether the graph contains the specified Vertex's `id`. The argument `id` is a String.

```python
class Graph:
    def __init__(self):
        self.vertices = {}
        
    def _create_vertex(self, id):
        return Vertex(id)
    
    def add_vertex(self, id):
        self.vertices[id] = self._create_vertex(id)
    
    def get_vertex(self, id):
        return self.vertices.get(id, None)
    
    def add_edge(self, start_v, end_v, weight=0):
        if start_v not in self.vertices:
            start = self.add_vertex(start_v)
        else:
            start = self.vertices[start_v]
        if end_v not in self.vertices:
            end = self.add_vertex(end_v)
        else:
            end = self.vertices[end_v]
        
        start.add_neighbour(end, weight)
            
    
    def get_neighbours(self, id):
        v = self.get_vertex(id)
        
        if v is not None:
            result = []
            for i in v.get_neighbours():
                result.append(i.id)
            return result
        else:
            return None
    
    def __contains__(self, val):
        return val in self.vertices.keys()
    
    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
            
    @property    
    def num_vertices(self):
        return len(self.vertices)
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
g = Graph()
assert g.vertices == {} and g.num_vertices == 0
g.add_vertex("A")
g.add_vertex("B")
g.add_vertex("C")
g.add_vertex("D")
g.add_vertex("E")
g.add_vertex("F")
assert g.num_vertices == 6
assert "A" in g
assert "B" in g
assert "C" in g
assert "D" in g
assert "E" in g
assert "F" in g
g.add_edge("A", "B")
g.add_edge("A", "C")
g.add_edge("B", "C")
g.add_edge("B", "D")
g.add_edge("C", "D")
g.add_edge("D", "C")
g.add_edge("E", "F")
g.add_edge("F", "C")
assert sorted(g.get_neighbours("A")) == ["B", "C"]
assert sorted(g.get_neighbours("B")) == ["C", "D"]
assert sorted(g.get_neighbours("C")) == ["D"]
assert [v.id for v in g] == ["A", "B", "C", "D", "E", "F"]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Create a subclass of `Vertex` called `VertexSearch`. This class has the following additional attributes:
- `colour`: which is a mark on the vertex during the search algorithm. It is of String data type and should be set to "white" by default.
- `d`: which is an Integer denoting the distance from other Vertex to the current Vertex in Breath-First-Search. This is also used to record discovery time in Depth-First-Search. This property should be initialized to `sys.maxsize` at the start.
- `f`: which is an Integer denoting the final time in Depth-First-Search. This property should be initialized to `sys.maxsize` at the start.
- `parent`: which is a reference to the parent Vertex object. This property should be set to `None` at the start.

```python
import sys

class VertexSearch(Vertex):
    def __init__(self, id=""):
        super().__init__(id)
        self.colour = "white"
        self.d = sys.maxsize
        self.f = sys.maxsize
        self.parent = None 
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
import sys

v = VertexSearch()
assert v.id == ""
assert v.colour == "white"
assert v.d == sys.maxsize
assert v.f == sys.maxsize
assert v.parent == None
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** **You should do this after you completed HW2.** Create a class `Search2D` which takes in an object `GraphSearch` for its initialization. The class should have the following methods:
- `clear_vertices()`: which sets the properties of all the vertices:
  - `colour` to "white"
  - `d` to `sys.maxsize`
  - `f` to `sys.maxsize`
  - `parent` to `None`.
 




```python
#Copy over GraphSearch from Homework
class GraphSearch(Graph):
    ##BEGIN SOLUTION
    def _create_vertex(self, id):
        return VertexSearch(id)
    ##END SOLUTION
    pass
```

```python
import sys

class Search2D:
    def __init__(self, g):
        self.graph = g
    
    def clear_vertices(self):
        for v in self.graph.vertices.values():
            v.colour = "white"
            v.d = sys.maxsize
            v.f = sys.maxsize
            v.parent = None
    
    def __iter__(self):
        return iter([v for v in self.graph])
    
    def __len__(self):
        return len([v for v in self.graph.vertices])
```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4 = Search2D(g4)
gs4.clear_vertices()

assert len(gs4) == 6
assert [v.id for v in gs4] == ["A", "B", "C", "D", "E", "F"]
assert [v.colour for v in gs4] == ["white" for v in range(len(gs4))]
assert [v.d for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.f for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.parent for v in gs4] == [None for v in range(len(gs4))]
```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Create a class `SearchBFS` which is a subclass of `Search2D`. This subclass should implement the Breadth First Search algorithm in the following methods:

- `search_from(start)`: which initializes the `d` and `parent` properties of each vertices in the graph from the `start` Vertex following Breadth-First-Search algorithm. Use your previous code that implements `Queue` data structure. 
- `get_shortest_path(start, dest)`: which returns a list of vertex ids that forms a shortest path from Vertex `start` to Vertex `dest`. Think how to solve this using recursion. This method should call `search_from()` if the distance at `start` Vertex is not zero. It should return a list containing ["No Path"] if there is no path from the start to the destination vertex. *Hint: you can make use another function for your recursion `get_path(start, dest, result)` where result is an empty list which will be populated in the recursive calls.*

```python
#Copy over Queue from Wk4 HMWK
class Queue:
    def __init__(self):
        self.__items = []
    
    def enqueue(self, item):
        self.__items.append(item)
    
    def dequeue(self):
        if self.__items == []:
            return None
        else:
            return self.__items.pop(0)
    
    def peek(self):
        return self.__items[0]
    
    @property
    def is_empty(self):
        return self.__items == []
    
    @property
    def size(self):
        return len(self.__items)
```

```python
class SearchBFS(Search2D):

    def search_from(self, start):
        self.clear_vertices()
        start = self.graph.get_vertex(start)
        if start is None:
            return
        start.colour = "grey"
        start.d = 0
        start.parent = None
        
        q = Queue()
        q.enqueue(start)
        
        while not q.is_empty:
            u = q.dequeue()
            for v in u.get_neighbours():
                if v.colour == "white":
                    v.colour = "grey"
                    v.d = u.d + 1
                    v.parent = u
                    q.enqueue(v)
            u.colour = "black"
    
    def get_shortest_path(self, start, dest):
        result = []
        self.get_path(start, dest, result)
        if result == []:
            return None
        else:
            return result
    
    def get_path(self, start, dest, result):
        u = self.graph.get_vertex(start)
        if u is None:
            return 
        else:
            v = self.graph.get_vertex(dest)
            if v is None:
                return
        
        if u.d != 0:
            self.search_from(start)
            
        if u == v:
            result.append(start)
        elif v.parent is None:
            result.append("No Path")
        else:
            self.get_path(start, v.parent.id, result)
            result.append(dest)
            

```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4 = SearchBFS(g4)

gs4.search_from("A")
assert gs4.graph.get_vertex("A").d == 0
assert gs4.graph.get_vertex("A").colour == "black"
assert gs4.graph.get_vertex("A").parent == None
assert gs4.graph.get_vertex("B").d == 1
assert gs4.graph.get_vertex("B").colour == "black"
assert gs4.graph.get_vertex("B").parent == gs4.graph.get_vertex("A")
assert gs4.graph.get_vertex("C").d == 1
assert gs4.graph.get_vertex("C").colour == "black"
assert gs4.graph.get_vertex("C").parent == gs4.graph.get_vertex("A")
assert gs4.graph.get_vertex("D").d == 2
assert gs4.graph.get_vertex("D").colour == "black"
gs4.graph.get_vertex("D").parent
#assert gs4.graph.get_vertex("D").parent == gs4.graph.get_vertex("B")
ans = gs4.get_shortest_path("A", "D")
assert ans == ["A", "B", "D"]
ans = gs4.get_shortest_path("E", "D")
assert ans == ["E", "F", "C", "D"]
```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```